Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Harsh Patel"
```

---

# <center>Major Assignment 5 Version 4</center>
-------------
<b><center>ENGINEER 1D04</center><b>
<b><center>Dr. Ashgar Bokhari</center><b>
<b><center>McMaster University, Fall 2017</center><b>

## Background
----------------
In physics, many kinematics problems can be simply modelled using a particle, or point mass, approximation. A one-dimensional particle with mass $m$ travels in only one direction, and has a position $r$, velocity $v$, and acceleration $a$ that change over time.

 In a numerical simulation of a particle, the flow of time is approximated by using intervals. For a period of time $t$ split into $n$ intervals, the change in time over each interval is given by $\Delta t = \frac{t}{n}$. The smaller $\Delta t$ is ($n$ is large), the better the approximation becomes. For a given time interval $\Delta t$, the position $r$ and velocity $v$ for a particle are updated as follows:
$$r_{\mathit{new}} = r_{\mathit{last}} + v_{\mathit{last}}\cdot\Delta t + \frac{1}{2}\cdot a \cdot \Delta t^2$$
$$v_{\mathit{new}} = v_{\mathit{last}} + a \cdot \Delta t$$

where $r_{\mathit{new}}$ and $v_{\mathit{new}}$ are the new (updated) position and velocity of the particle, and $r_{\mathit{last}}$ and $v_{\mathit{last}}$ are the position and velocity during the last interval.

This program will simulate a particle being fired from a cannon into the air, then being pulled back to the ground by gravity. A mass $m$ will experience an initial energy input $E$, setting its velocity according to the following equation:
$$v = \sqrt{\frac{2\cdot E}{m}}$$
                                    
Design, implement, and test a program that satisfies the requirements below.

-------------------
## Requirements
1. The library includes a Python class named `Particle1D`. (Note: **1** in Particle**1**D is the number one) <br><br>

2. The `__init__` method for `Particle1D` takes first formal parameter `self` and four numbers $m$, $r_0$, $v_0$, and $a_0$ as input **in that order**, and stores each of them as a class instance variable. $m$, $r$, $v$ and $a$ represent the mass, starting position, starting velocity, and starting acceleration of the particle, respectively.<br><br>

3. The class contains accessors for $m$, $r$, $v$, and $a$. The names of the accessors are `getM()`, `getR()`, `getV()` and `getA()`, respectively. <br><br>

4. The class contains a mutator `update(self, delta_t)` that has a first formal parameter `self` and a normal parameter time interval $\Delta t$ , and updates $r$ and $v$ according to the formulas above. Be careful not to use $v_{\mathit{new}}$ in the calculation for $r_{\mathit{new}}$.<br><br>

5. The class contains a method `isOnGround(self)`, that has first formal parameter `self` and returns `True` if $r \le 0$, and `False` otherwise.<br><br>

6.  The class contains a mutator `fireCannon(self, energy):` that has a first formal parameter `self` and a normal parameter $energy$, and sets $v$ according to the energy formula above. Note that $v$ is re-calculated directly from $E$, and $m$.<br><br>

7. **Outside** the class there should be a function `runCannonSim(delta_t, energy, gravity, mass)` that takes a time interval $\Delta t$, a float `energy`, a float `gravity`, and a float `mass` as input, and simulates a mass being fired from a cannon into the air. First, a `Particle1D` object is created with mass `mass`, initial position `0`, initial velocity `0`, and initial acceleration `gravity`. The particle is then fired from the cannon using `energy` and `fireCannon`. The simulation will run until the particle hits the ground ($r\le0$). Each "cycle" of the simulation recalculates the position and velocity of the particle using $\Delta t$ and `update`. The function should return a list containing the starting point of the particle, and its new position ($r$) at the end of each cycle. For example, if the particle has a `mass` = 10, `energy` = 1000 and `delta_t` = 0.4, the returned list will look something like: `[0.0, 8.177708498984762, 9.914562748477142, 10.083416997969522, 8.6842712474619, 5.71712549695428, 1.1819797464466593, -4.921166004060961]`.<br><br>

8. The program requires very little besides the function definitions. There is no main().<br><br>

9. Your name, MacID, student number, and the date are given in comments at the top of the first Python cell below.<br><br>

10. Your answers to the design questions and test plan are given in the appropriate Markdown cells below.<br><br>

11. Your program MUST have valid Python syntax and it must run without errors. Ensure that your program runs properly by running it before you submit.<br><br>

12. You must sign out with a TA or IAI after you have submitted your lab at the submission station. Failure to do so could result in a mark of zero.
------------------
## Design and Implementation Instructions

1. You may assume that inputs will be within a valid range and of the correct type, and that exception handling is not necessary.<br><br>
2. Follow the function syntax **EXACTLY** as given, including spelling, capitalization and the order of function paramaters.

------------------

## Particle1D Class with runSpringSim(p, cycles, delta_t, k) function - (10 marks)

<!-- grade_id: particle1D, points: 0, grade: false, solution: true -->
```python
class Particle1D:
    
    def __init__(self,m,r0,v0,a0):#constructor
        self.m=m
        self.r=r0
        self.v=v0
        self.a=a0
    
    #accessors
    def getM(self):
        return self.m
    def getR(self):
        return self.r
    def getV(self):
        return self.v
    def getA(self):
        return self.a
    
    #mutators
    def update(self,dt):#changes both r and v values
        self.r = self.r + self.v*dt + 0.5*self.a*(dt**2)
        self.v = self.v + self.a*dt
        
    def isOnGround(self):# r<=0 is already an expression that would equate to True or False, so I just returned that expression
        return self.r<=0
    
    def fireCannon(self,energy):#calculates new v
        self.v = ((2*energy)/self.m)**0.5
        

def runCannonSim(dt,energy,gravity,mass):
    
    rlist=[0.0]#creating a list that will later be returned with initial position 0
    p1 = Particle1D(mass,0,0,gravity)#creating a new particle
    p1.fireCannon(energy)#updating new velocity with fire cannon
    p1.update(dt)#updating the first time after fireCannon
    
    while not(p1.isOnGround()):#while loop to repeat until it lands on ground
        p1.update(dt)
        rlist.append(p1.getR())#add the updated r value to list
    
    return rlist
```

<!-- grade_id: particle1D_tests, points: 10, grade: true, solution: false -->
```python
test_particle_1 = Particle1D(1, 0, 2, 0)
test_particle_2 = Particle1D(2.5, -1.5, -1.5, 2)
test_particle_3 = Particle1D(4,1,1,2)
test_particle_4 = Particle1D(1,-1,0,0)

test_particle_1.update(0.1)
test_particle_2.update(0.2)
assert test_particle_1.getM() == 1
assert test_particle_1.getR() == 0.2
assert test_particle_1.getV() == 2
assert test_particle_1.getA() == 0
assert test_particle_2.getM() == 2.5
assert test_particle_2.getR() == -1.76
assert test_particle_2.getV() == -1.1
assert test_particle_2.getA() == 2.0

assert test_particle_4.isOnGround() == True

test_particle_3.fireCannon(18)
assert test_particle_3.getM() == 4
assert test_particle_3.getR() == 1
assert test_particle_3.getV() == 3
assert test_particle_3.getA() == 2

assert runCannonSim(0.4, 1000, -9.8, 10) == [0.0, 8.177708498984762, 9.914562748477142, 10.083416997969522, 8.6842712474619, 5.71712549695428, 1.1819797464466593, -4.921166004060961]
assert runCannonSim(0.6, 400, -4.9, 5) == [0.0, 11.650932768808222, 14.830399153212333, 16.245865537616442, 15.897331922020554, 13.784798306424666, 9.908264690828778, 4.2677310752328905, -3.136802540362997]
assert runCannonSim(0.1, 0, -3, 2) == [0.0]
print("Tests Successful")
```

Output:
```
Tests Successful
```

## Code Legibility (3 Points)
----------
Your code will be marked on commenting and code legibility.<br\>
The mark breakdown is as follows:<br\>
1 mark for using appropriate variable names that indicate what is being stored in that variable<br\>
1 mark for leaving comments on major parts of your code such as where you read the file or calculate a summation<br\>
1 mark for general legibility. The TA's should be able to understand your code without spending hours reading it. For example do not put your code in one very long line as this is hard for someone else reading your code to understand

<h6>Ignore the empty cell below. You do NOT need to type anything in it </h6>

<!-- grade_id: cell-921fcfa085debad8, points: 3, grade: true, solution: true -->
jsh

## Design Question (4 Points)
--------------
1. What is represented by the list of values that `runCannonSim` outputs? What would be a good use for this list?
2. Is it possible to create an object of class `Particle1D` inside a method of class `Particle1D`?
3. What is the difference between an accessor and a constructor?
4. When calling `fireCannon`, is it necessary to pass in the parameter `self`?

Enter your answers into the Markdown cell below.

<!-- grade_id: major_5_v4_dq_answers, points: 4, grade: true, solution: true -->
1. runCannonSim outputs a list of the positions of the particle being fired for every time interval, dt, until it hits the ground. A good use for this list would be graphing. We can plot the position of the particle vs. time by using the matplotlib library, since we already know dt and can create a list corresponding to that. The x-list would be the time values, and the y-list would be the list returned by runCannonSim.

2. Yes it is. You can call the constructor of a class inside a method of the class itself.

3. An accessor just returns a value such as a class instance variable. On the other hand, the constructor generally defines the class instance variables, giving them their initial values. The constructor also cannot return anything other than None, whereas accessors can.

4. No it is not, self is already implied when calling a method in a class. It is a formal parameter and does not have to be accounted for when calling a class method.

## Testing Plan (3 Points)
---------------
Produce a test plan in the Markdown cell below, in the following form:

`Test i for function j
Input: inputs for function j
Expected Output: expected output for function j
Actual Output: actual output for function j
Result: Pass/Fail
`

Note: The actual output should be what the program produces, even if your output does not match the expected output.

You must have $NO$ $LESS$ $THAN$ $3$ $TEST$ $CASES$. Have at least 1 case where your program does not output an error. For the other cases, we encourage you to try and find test cases where your program would output an error (not mandatory, just recommended). That is, where the expected output is an error.

<!-- grade_id: major_5_v4_test_plan_answers, points: 3, grade: true, solution: true -->
Test cases 1-4 are for an object: P1 = Particle1D(4,2,6,1)

1.
Test 1 for function getM(self)
Input: None
Expected Output: 4
Actual Output: 4
Result: Pass

2.
Test 1 for update(self,dt)
Input: update(2)
Expected Output: None, it is a mutator and does not return anything, however the corresponding R and V values should be 16, and
                 8 respectively
Actual Output: None, getR() was 16.0 and getV() was 8
Result: Pass

3.
Test 2 for update(self,dt)
Input: update([2])
Expected Output: TypeError, since it expects an integer or float value in the calculations for Rnew and Vnew
Actual Output: TypeError: unsupported operand type(s) for \*\* or pow(): 'list' and 'int'
Result: Pass

4.
Test 1 for isOnGround(self)
Input: isOnGround(4)
Expected Output: Error, since the isOnGround(self) function does not accept any normal parameters
Actual Output: TypeError: isOnGround() takes 1 positional argument but 2 were given
Result: Pass

5.
Test 1 for runCannonSim(dt,energy,gravity,mass)
Input: runCannonSim(2,3,6,-4)
Expected Output: Some mathematical error, since mass is used in a square root calculation and therefore cannot be negative
Actual Output: TypeError: unorderable types: complex() <= int()
Result: Pass    (The type of error was not correct but the logic in why the error occured was correct)